# Linear Regression From Scratch
In this notebook you will implement linear regression using the least squares method.

The notebook will guide you in the general steps. You may use only numpy, pandas, and matplotlib for this exercise.

#### Math Background
The task is to solve linear regression for the data set of ```x``` and ```y```. That is, fit a line y to the data such that ```y = a + bx + e```. Where a and b are coefficents and e is an error term.

We know that ```b = SUM ( xi - x_mean) * (yi - y_mean) / SUM (xi - x_mean)^2``` where ```xi```, and ```yi``` are the indivdual observations and ```x_mean```, ```y_mean``` are means of all ```xi``` and ```yi```.

So a reasonable solution then for a fit is ```a = y_mean - b * x_mean```.


#### Implementation Steps
1. Load the file ```reg_data.csv```
2. Display the data with a scatter plot. Write a markdown cell and comment on the structure of the data and what kind of line you think will fit the data. 
3. Implement a function of the least squares method.
4. Plot the predictions from your function.
5. Comment on the output of your function. Does it make sense? Can it be made better? If yes, how?

```python
#import the minimum packages
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split
from sklearn import datasets, linear_model
from sklearn.linear_model import LinearRegression
```

<!-- grade_id: cell-4104816f50e4844b, points: 0, grade: false, solution: true -->
```python
#load and inspect the data
df = pd.read_csv('reg_data.csv')
df.head()

```

<!-- grade_id: cell-3676b87d8fc79bbd, points: 1, grade: true, solution: false -->
```python
assert df.shape == (100,2)
```

<!-- grade_id: cell-7f576e50392bc5e9, points: 0, grade: false, solution: true -->
```python
#plot the data (scatter)
ax1 = df.plot.scatter(x='X',
                      y='Y',
                      c='DarkBlue')
```

<!-- grade_id: cell-54891567ed3179cb, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-161d857ab77810d3, points: 1, grade: true, solution: false -->
```python
assert x_train.shape == (80,)
assert x_test.shape == (20,)
assert y_train.shape == (80,)
assert y_test.shape == (20,)
```

```python
df.describe()
```

Output:
```
X           Y
count  100.000000  100.000000
mean    48.958341   72.735051
std      9.746379   16.658187
min     25.128485   31.707006
25%     41.526164   60.788586
50%     49.634970   72.179542
75%     56.762631   83.215080
max     70.346076  118.591217
```

<!-- grade_id: cell-d5e750d407949b20, points: 0, grade: false, solution: true -->
```python
#implement a least squares function to find a, b

x_mean = np.mean(df['X'])
y_mean = np.mean(df['Y'])
b_upper = np.max(df['Y'])
b_lower = np.min(df['Y'])
print(b)
print(a)



```

```python
# YOUR CODE HERE
X = df['X'].values.reshape(-1,1)
Y = df['Y'].values.reshape(-1,1)
x_train, x_test, y_train, y_test = train_test_split(X, Y, test_size=0.2)
regressor = LinearRegression()
regressor.fit(x_train, y_train)
print(regressor.coef_)
print(regressor.intercept_)

```

Output:
```
[[1.36990554]]
[5.39982785]
```

```python
a = float(regressor.coef_[0])
b = float(regressor.intercept_[0])
print(b)
print(a)


```

Output:
```
5.399827852237422
1.3699055352233545
```

<!-- grade_id: cell-c7fea6e45dc92701, points: 5, grade: true, solution: false -->
```python
assert abs(a - 7.72) <=0.03
assert abs(b - 1.32) <=0.03
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-46-94ffca13c1d9> in <module>
----> 1 assert abs(a - 7.72) <=0.03
      2 assert abs(b - 1.32) <=0.03

AssertionError:
```

<!-- grade_id: cell-721384e1a2152077, points: 0, grade: false, solution: true -->
```python
line = 0
x = np.array([])
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-71f1201b79ff7e44, points: 5, grade: true, solution: false -->
```python
assert abs(a +3*b - 11.69) <= 0.05
assert abs(a +100*b - 140) <= 1
```

<!-- grade_id: cell-e8d9edf92254aa00, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8b80d3f35ea3b407, points: 0, grade: false, solution: true -->
```python
#Classify your test data in to classes
#if the Y value for a certain X is lower than the line then the class is 0
class_0 = []
class_1 = []

# YOUR CODE HERE
raise NotImplementedError()
        
class_0 = np.array(class_0)
class_1 = np.array(class_1)
print(class_0.shape)
print(class_1.shape)
```

<!-- grade_id: cell-73e31e3494663d82, points: 5, grade: true, solution: false -->
```python
assert 9 < class_0.shape[0] < 13
assert 7 < class_1.shape[0] < 11
```

<!-- grade_id: cell-db81fd9fa3a21bdf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```